---
title       : Insert the chapter title here
description : Insert the chapter description here
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf




--- type:NormalExercise lang:r xp:100 skills:1 key:0e24f71c91
## Introduction to Advanced Cool stuff

Some things here

*** =instructions
- And here

*** =hint
- Hintje

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}
# Some sample code
```

*** =solution
```{r}
# Some solution code
```

*** =sct
```{r}

success_msg("Good work!")
```

